# Cybersecurity
**CS1302 Introduction to Computer Programming**
___

Python is a popular tool among hackers and engineers. In this lab, you will learn Cryptology in cybersecurity, which covers
- [Cryptography](https://en.wikipedia.org/wiki/Cryptography): Encryption and decryption using a cipher.
- [Cryptanalysis](https://en.wikipedia.org/wiki/Cryptanalysis): Devising an attack to break a cipher.

## Caesar symmetric key cipher

We first implements a simple cipher called the [Caesar cipher](https://en.wikipedia.org/wiki/Caesar_cipher).

```python
%%html
<iframe width="912" height="513" src="https://www.youtube.com/embed/sMOZf4GN3oc" frameborder="0" allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture" allowfullscreen></iframe>
```

Output:
```
<IPython.core.display.HTML object>
```

### Encrypt/decrypt a character

**How to encrypt a character?**

The following code encrypts a character `char` using a non-negative integer `key`.

```python
cc_n = 1114112


def cc_encrypt_character(char, key):
    '''
    Return the encryption of a character by an integer key using Caesar cipher.
    
    Parameters
    ----------
    char (str): a unicode (UTF-8) character to be encrypted.
    key (int): secret key to encrypt char.
    '''
    char_code = ord(char)
    shifted_char_code = (char_code + key) % cc_n
    encrypted_char = chr(shifted_char_code)
    return encrypted_char
```

For example, to encrypt the letter `'A'` using a secret key `5`:

```python
cc_encrypt_character('A', 5)
```

Output:
```
'F'
```

```python
cc_n = 1114112
# cc_n = 11141120 #ValueError: chr() arg not in range(0x110000) which is equal to 1114112

key = 5


plain_char = 'A'
print('plain_char:', plain_char)


char_code = ord(plain_char)
print('char_code:', char_code)


print('key:', key)

char_code_plus_key = (char_code + key)
print('char_code_plus_key:', char_code_plus_key)

shifted_char_code = (char_code + key) % cc_n
print('shifted_char_code:', shifted_char_code)

encrypted_char = chr(shifted_char_code)
print('encrypted_char:', encrypted_char)

print(plain_char, ' encrypts to ', encrypted_char, '(key:', key, ')')


```

Output:
```
plain_char: A
char_code: 65
key: 5
char_code_plus_key: 70
shifted_char_code: 70
encrypted_char: F
A  encrypts to  F (key: 5 )
```

The character `'A'` is encrypted to the character `'F'` as follows:

1. `ord(char)` return the integer `65` that is the code point (integer representation) of the unicode of `'A'`. 
2. `(char_code + key) % cc_n` cyclic shifts the code by the key `5`.
3. `chr(shifted_char_code)` converts the shifted code back to a character, which is `'F'`.

| Encryption                      |     |       |     |     |     |     |     |     |
| ------------------------------- | --- | ----- | --- | --- | --- | --- | --- | --- |
| `char`                          | ... | **A** | B   | C   | D   | E   | F   | ... |
| `ord(char)`                     | ... | **65**| 66  | 67  | 68  | 69  | 70  | ... |
| `(ord(char) + key) % cc_n`      | ... | **70**| 71  | 72  | 73  | 74  | 75  | ... |
| `(chr(ord(char) + key) % cc_n)` | ... | **F** | G   | H   | I   | J   | K   | ... |

You may learn more about `ord` and `chr` from their docstrings:

```python
help(ord)
help(chr)
```

Output:
```
Help on built-in function ord in module builtins:

ord(c, /)
    Return the Unicode code point for a one-character string.

Help on built-in function chr in module builtins:

chr(i, /)
    Return a Unicode string of one character with ordinal i; 0 <= i <= 0x10ffff.
```

**How to decrypt a character?**

Mathematically, we define the encryption and decryption of a character for Caesar cipher as
$$ \begin{aligned} E(x,k) &:= x + k \mod n & \text{(encryption)} \\
D(x,k) &:= x - k \mod n & \text{(decryption),} \end{aligned}
$$
where $x$ is the character code in $\{0,\dots,n\}$ and $k$ is the secret key. `mod` operator above is the modulo operator. In Mathematics, it has a lower precedence than addition and multiplication and is typeset with an extra space accordingly.

The encryption and decryption satisfies the recoverability condition
$$ D(E(x,k),k) = x $$
so two people with a common secret key can encrypt and decrypt a character, but others not knowing the key cannot. This is a defining property of a [symmetric cipher](https://en.wikipedia.org/wiki/Symmetric-key_algorithm).  

The following code decrypts a character using a key.

```python
def cc_decrypt_character(char, key):
    '''
    Return the decryption of a character by the key using Caesar cipher.
    
    Parameters
    ----------
    char (str): a unicode (UTF-8) character to be decrypted.
    key (int): secret key to decrypt char.
    '''
    char_code = ord(char)
    shifted_char_code = (char_code - key) % cc_n
    decrypted_char = chr(shifted_char_code)
    return decrypted_char
```

For instance, to decrypt the letter `'F'` by the secret key `5`:

```python
cc_decrypt_character('F',5)
```

Output:
```
'A'
```

```python
cc_n = 1114112
# cc_n = 11141120 #ValueError: chr() arg not in range(0x110000) which is equal to 1114112

key = 5


ciphertext = 'F'
print('ciphertext:', ciphertext)


char_code = ord(ciphertext)
print('char_code:', char_code)


print('key:', key)

char_code_minus_key = (char_code - key)
print('char_code_minus_key:', char_code_minus_key)

shifted_char_code = (char_code - key) % cc_n
print('shifted_char_code:', shifted_char_code)

decrypted_char = chr(shifted_char_code)
print('encrypted_char:', decrypted_char)

print(ciphertext, ' decrypts to ', decrypted_char, '(key:', key, ')')
```

Output:
```
ciphertext: F
char_code: 70
key: 5
char_code_minus_key: 65
shifted_char_code: 65
encrypted_char: A
F  decrypts to  A (key: 5 )
```

The character `'F'` is decrypted back to `'A'` because
`(char_code - key) % cc_n` reverse cyclic shifts the code by the key `5`.

| Encryption                      |     |       |     |     |     |     |     |     | Decryption                      |
| ------------------------------- | --- | ----- | --- | --- | --- | --- | --- | --- | ------------------------------- |
| `char`                          | ... | **A** | B   | C   | D   | E   | F   | ... | `(chr(ord(char) - key) % cc_n)` |
| `ord(char)`                     | ... | **65**| 66  | 67  | 68  | 69  | 70  | ... | `(ord(char) - key) % cc_n`      |
| `(ord(char) + key) % cc_n`      | ... | **70**| 71  | 72  | 73  | 74  | 75  | ... | `ord(char)`                     |
| `(chr(ord(char) + key) % cc_n)` | ... | **F** | G   | H   | I   | J   | K   | ... | `char`                          |

**Exercise** Why did we set `cc_n = 1114112`? Explain whether the recoverability property may fail if we set `cc_n` to a bigger number or remove `% cc_n` for both `cc_encrypt_character` and `cc_decrypt_character`.

<!-- grade_id: modulo, points: 1, grade: true, solution: true -->
**answer:**  cc_n is set to 1114112 because that is the largest character that can be coded with Unicode. No matter how large the key is, we want to keep the encrypted character with in the unicode range by using the modulo 1114112.  

**i learnt this after hard work**

### Encrypt a plaintext and decrypt a ciphertext

Of course, it is more interesting to encrypt a string instead of a character. The following code implements this in one line.

```python
def cc_encrypt(plaintext, key):
    '''
    Return the ciphertext of a plaintext by the key using Caesar cipher.
    
    Parameters
    ----------
    plaintext (str): a unicode (UTF-8) message in to be encrypted.
    key (int): secret key to encrypt plaintext.
    '''
    return ''.join([chr((ord(char) + key) % cc_n) for char in plaintext])
```

The above function encrypts a message, referred to as the *plaintext*, by replacing each character with its encryption.  
This is referred to as a [*substitution cipher*](https://en.wikipedia.org/wiki/Substitution_cipher).

**Exercise** Define a function `cc_decrypt` that
- takes a string `ciphertext` and an integer `key`, and
- returns the plaintext that encrypts to `ciphertext` by the key using Caesar cipher.

<!-- grade_id: cc_decrypt, points: 0, grade: false, solution: true -->
```python
def cc_decrypt(ciphertext, key):
    '''
    Return the plaintext that encrypts to ciphertext by the key using Caesar cipher.
    
    Parameters
    ----------
    ciphertext (str): message to be decrypted.
    key (int): secret key to decrypt the ciphertext.
    '''
    return ''.join([chr((ord(char) - key) % cc_n) for char in ciphertext])
```

<!-- grade_id: test-cc_decrypt, points: 1, grade: true, solution: false -->
```python
# tests
assert cc_decrypt(r'bcdefghijklmnopqrstuvwxyz{',1) == 'abcdefghijklmnopqrstuvwxyz'
assert cc_decrypt(r'Mjqqt1%\twqi&',5) == 'Hello, World!'
```

#### My Tests

```python
print(cc_encrypt('It is a fucking waste of time for people without any imagination', 1))
```

Output:
```
Ju!jt!b!gvdljoh!xbtuf!pg!ujnf!gps!qfpqmf!xjuipvu!boz!jnbhjobujpo
```

```python
print(cc_encrypt('I see something in you, but i don\'t know what it is', 1))
```

Output:
```
J!tff!tpnfuijoh!jo!zpv-!cvu!j!epo(u!lopx!xibu!ju!jt
```

```python
print(cc_encrypt('ABCDEFG𩪂ജ', 5))
print(cc_encrypt('Fiaz Idris', 1))
print(cc_encrypt('Fuck You!', 8))
print(cc_encrypt('Fuck You!', 9))

```

Output:
```
FGHIJKL𩪇ഡ
Gjb{!Jesjt
N}ks(aw})
O~lt)bx~*
```

```python
print(cc_decrypt('FGHIJKL𩪇ഡ', 5))
print(cc_decrypt('Gjb{!Jesjt', 1))
print(cc_decrypt('N}ks(aw})', 8))
print(cc_decrypt('O~lt)bx~*', 9))
```

Output:
```
ABCDEFG𩪂ജ
Fiaz Idris
Fuck You!
Fuck You!
```

## Brute-force attack

### Create an English dictionary

You will launch a brute-force attack to guess the key that encrypts an English text. The idea is simple: 

- You try decrypting the ciphertext with different keys, and 
- see which of the resulting plaintexts make most sense (most english-like).

To check whether a plaintext is English-like, we need to have a list of English words. One way is to type them out
but this is tedious. Alternatively, we can obtain the list from the *Natural Language Toolkit (NLTK)*: 

```python
import nltk
nltk.download('words')
from nltk.corpus import words
```

Output:
```
[nltk_data] Downloading package words to
[nltk_data]     C:\Users\jupyter_user\AppData\Roaming\nltk_data...
[nltk_data]   Package words is already up-to-date!
```

`words.words()` returns a list of words. We can check whether a string is in the list using the operator `in`. 

```python
type(words.words())
```

Output:
```
list
```

```python
len(words.words())
```

Output:
```
236736
```

```python
import pandas as pd
dfwords = pd.DataFrame(words.words())
dfwords.head(20)
```

Output:
```
0
0           A
1           a
2          aa
3         aal
4       aalii
5         aam
6        Aani
7    aardvark
8    aardwolf
9       Aaron
10    Aaronic
11  Aaronical
12   Aaronite
13  Aaronitic
14       Aaru
15         Ab
16        aba
17    Ababdeh
18     Ababua
19       abac
```

```python
dfwords.sample(20)
```

Output:
```
0
96778        intersession
21776                bhoy
9480           antetemple
107172        lipothymial
158230  provincialization
34377            chevance
189280           stepaunt
80697             Guetare
41614        consensually
119814             mosser
76089              gelada
94827           injection
210700        unabundance
214919       undesignated
217858         unidextral
26688        bronchoscopy
158048         protostele
144344         pheasantry
117613    mistrustfulness
11907      appositionally
```

```python
dfwords['lowerword'] = dfwords[0].str.lower()

```

```python
dfwords.tail()
```

Output:
```
0  lowerword
236731     yellow     yellow
236732        yes        yes
236733  yesterday  yesterday
236734        you        you
236735      young      young
```

```python
for word in 'Ada', 'ada', 'Hello', 'hello':
    print('{!r} in dictionary? {}'.format(word, word in words.words()))
```

Output:
```
'Ada' in dictionary? True
'ada' in dictionary? False
'Hello' in dictionary? False
'hello' in dictionary? True
```

However there are two issues:
- Checking membership is slow for a long list.
- Both 'Hello' and 'ada' are English-like but they are not in the words_list.

**Exercise** Using the method `lower` of `str` and the constructor `set`, assign `dictionary` to a set of lowercase English words from `words.words()`.

<!-- grade_id: nltk, points: 0, grade: false, solution: true -->
```python
#Great Uncle's greatest tricks
wordlist = words.words() # just a word list variable with duplicates
dictionary = {word.lower() for word in wordlist} # it makes a set that means unique.. we call this set comprehension

```

```python
wordlist[:10]
```

Output:
```
['A',
 'a',
 'aa',
 'aal',
 'aalii',
 'aam',
 'Aani',
 'aardvark',
 'aardwolf',
 'Aaron']
```

```python
sorted(dictionary)[:10]
```

Output:
```
['a',
 'aa',
 'aal',
 'aalii',
 'aam',
 'aani',
 'aardvark',
 'aardwolf',
 'aaron',
 'aaronic']
```

```python
print(len(wordlist))
print(len(dictionary))
```

Output:
```
236736
234377
```

<!-- grade_id: test-nltk, points: 1, grade: true, solution: false -->
```python
# tests
assert isinstance(dictionary,set) and len(dictionary) == 234377
assert all(word in dictionary for word in ('ada', 'hello'))
assert all(word not in dictionary for word in ('Ada', 'hola'))
### BEGIN TESTS
assert 'world' in dictionary
assert not 'mundo' in dictionary
### END TESTS
```

### Identify English-like text

To determine how English-like a text is, we calculate the following score:
$$
\frac{\text{number of English words in the text}}{\text{number of tokens in the text}} 
$$
where tokens are substrings (not necessarily an English word) separated by white space characters in the text.

```python
def tokenizer(text):
    '''Returns the list of tokens of the text.'''
    return text.split()

def get_score(text):
    '''Return the fraction of tokens which appear in dictionary.'''
    tokens = tokenizer(text)
    words = [token for token in tokens if token in dictionary]
    return len(words)/len(tokens)

# tests
get_score('hello world'), get_score('Hello, World!')
```

Output:
```
(1.0, 0.0)
```

As shown in tests above, the code fails to handle text with punctuations and uppercase letters properly.  
In particular, 
- while `get_score` recognizes `hello world` as English-like and returns the maximum score of 1, 
- it fails to recognize `Hello, World!` as English-like and returns the minimum score of 0.

Why? This is because every words in `dictionary`
- are in lowercase, and
- have no leading/trailing punctuations.

**Exercise** Define a funtion `tokenizer` that 
- takes a string `text` as an argument, and
- returns a `list` of tokens obtained by
  1. splitting `text` into a list using `split()`;
  2. removing leading/trailing punctuations in `string.punctuation` using the `strip` method; and
  3. converting all items of the list to lowercase using `lower()`.

```python
import string # import string as we need to use string.punctuation
string.punctuation
```

Output:
```
'!"#$%&\'()*+,-./:;<=>?@[\\]^_`{|}~'
```

```python
mystring = 'Hello, World!'

```

```python
mystring.split()
```

Output:
```
['Hello,', 'World!']
```

```python
[eachword.strip(string.punctuation).lower() for eachword in mystring.split()]
```

Output:
```
['hello', 'world']
```

<!-- grade_id: tokenizer, points: 0, grade: false, solution: true -->
```python
# another of your Great Uncle's tricks.. but don't worry.. this is baby stuff for someone above me.... which in reality is most people
# in this domain.
# strip only removes leading and trailing chracters, not in between

def tokenizer(text):
    '''Returns the list of tokens of the text such that 
    1) each token has no leading or training spaces/punctuations, and 
    2) all letters in each tokens are in lowercase.'''
    return [eachword.strip(string.punctuation).lower() for eachword in text.split()]
```

```python
tokenizer(mystring)
```

Output:
```
['hello', 'world']
```

```python
'a,b'.strip(string.punctuation)
```

Output:
```
'a,b'
```

```python
'Jean-Pierre?'.strip(string.punctuation).lower()
```

Output:
```
'jean-pierre'
```

<!-- grade_id: test-tokenizer, points: 1, grade: true, solution: false -->
```python
# tests
assert tokenizer('Hello, World!') == ['hello', 'world']
assert get_score('Hello, World!') >= 0.99999
assert tokenizer('Do you know Jean-Pierre?') == ['do', 'you', 'know', 'jean-pierre']
assert get_score('Do you know Jean-Pierre?') >= 0.99999
```

### Launch a brute-force attack

**Exercise** Define the function `cc_attack` that 
- takes as arguments
    - a string `ciphertext`,
    - a floating point number `threshold` in the interval $(0,1)$ with a default value of $0.6$, and
- returns a generator that  
    - generates one-by-one in ascending order guesses of the key that
    - decrypt `ciphertext` to texts with scores at least the `threshold`.

```python
import random
```

<!-- grade_id: cc_attack, points: 0, grade: false, solution: true -->
```python
def cc_attack(ciphertext, threshold = 0.6):
    '''Returns a generator that generates the next guess of the key that 
    decrypts the ciphertext to a text with get_score(text) at least the threshold.
    '''
    for i in range(cc_n):
        decrypted_text = cc_decrypt(ciphertext, i )
        decrypted_text_score = get_score(decrypted_text)
#         print(decrypted_text)
        if decrypted_text_score >= threshold:
            print('in cc_attack:', i, decrypted_text)
            yield i
    
    # YOUR CODE HERE
    pass
```

```python
original_text = 'lord of the world'
cc_n = 1114112
key = random.randint(0, cc_n)
key = cc_n - 1


encrypted_text = cc_encrypt(original_text, key)
print(original_text, ' >> ', encrypted_text)

print('attacking ', encrypted_text)

key_generator = cc_attack(encrypted_text)
key_guess = next(key_generator)

text = cc_decrypt(encrypted_text, key_guess)
print('guess of the key: {}\nscore: {}\ntext :{}'.format(key_guess,get_score(text),text))

print('================')
print(encrypted_text, ' decrypted to ', text, ' using key: ', key_guess, ' with score: ', get_score(text) )
```

Output:
```
lord of the world  >>  knqcnesgdvnqkc
attacking  knqcnesgdvnqkc
in cc_attack: 1114111 lord of the world
guess of the key: 1114111
score: 1.0
text :lord of the world
================
knqcnesgdvnqkc  decrypted to  lord of the world  using key:  1114111  with score:  1.0
```

```python
encrypted_text = 'めゞゥゥエづすゐエカゥゝず'

key_generator = cc_attack(encrypted_text)
key_guess = next(key_generator)

text = cc_decrypt(encrypted_text, key_guess)
print('guess of the key: {}\nscore: {}\ntext :{}'.format(key_guess,get_score(text),text))

print('================')
print(encrypted_text, ' decrypted to ', text, ' using key: ', key_guess, ' with score: ', get_score(text) )
```

Output:
```
in cc_attack: 12345 Hello, World!
guess of the key: 12345
score: 1.0
text :Hello, World!
================
めゞゥゥエづすゐエカゥゝず  decrypted to  Hello, World!  using key:  12345  with score:  1.0
```

```python
cc_encrypt("Hello, World!",12345)
```

Output:
```
'めゞゥゥエづすゐエカゥゝず'
```

```python
cc_decrypt('めゞゥゥエづすゐエカゥゝず', 12345)
```

Output:
```
'Hello, World!'
```

<!-- grade_id: test-cc_attack, points: 1, grade: true, solution: false -->
```python
# tests
ciphertext = cc_encrypt("Hello, World!",12345)
key_generator = cc_attack(ciphertext)
key_guess = next(key_generator)
assert key_guess == 12345
text = cc_decrypt(ciphertext, key_guess)
print('guess of the key: {}\nscore: {}\ntext :{}'.format(key_guess,get_score(text),text))
```

Output:
```
12345 Hello, World!
guess of the key: 12345
score: 1.0
text :Hello, World!
```

## Challenge

**I just learnt this from the wikipedia... and it can be done with some hard work!**  
I will do it later for my own interest... but do not want to spend time now!

Another symmetric key cipher is [columnar transposition cipher](https://en.wikipedia.org/wiki/Transposition_cipher#Columnar_transposition). A transposition cipher encrypts a text by permuting instead of substituting characters.

**Exercise** Study and implement the irregular case of the [columnar transposition cipher](https://en.wikipedia.org/wiki/Transposition_cipher#Columnar_transposition) as described in Wikipedia page. Define the functions 
- `ct_encrypt(plaintext, key)` for encryption, and 
- `ct_decrypt(ciphertext, key)` for decryption. 

You can assume the plaintext is in uppercase and has no spaces/punctuations. 

*Hints:* See the text cases for an example of `plaintext`, `key`, and the corresponding `ciphertext`. You can but are not required to follow the solution template below:

```Python
def argsort(seq):
    '''A helper function that returns the tuple of indices that would sort the
    sequence seq.'''
    return tuple(x[0] for x in sorted(enumerate(seq), key=lambda x: x[1]))


def ct_idx(length, key):
    '''A helper function that returns the tuple of indices that would permute 
    the letters of a message according to the key using the irregular case of 
    columnar transposition cipher.'''
    last_row_size = length % len(key)
    num_rows = length // len(key) + (1 if last_row_size else 0)
    seq = tuple(range(length))
    matrix = ____________________________________________________________________
    return tuple(matrix[i][j] for j in argsort(key)
                 for i in range(num_rows - (1 if j >= last_row_size else 0)))


def ct_encrypt(plaintext, key):
    '''
    Return the ciphertext of a plaintext by the key using the irregular case
    of columnar transposition cipher.
    
    Parameters
    ----------
    plaintext (str): a message in uppercase without punctuations/spaces.
    key (str): secret key to encrypt plaintext.
    '''
    return ''.join([plaintext[i] for i in ct_idx(len(plaintext), key)])


def ct_decrypt(ciphertext, key):
    '''
    Return the plaintext of the ciphertext by the key using the irregular case
    of columnar transposition cipher.
    
    Parameters
    ----------
    ciphertext (str): a string in uppercase without punctuations/spaces.
    key (str): secret key to decrypt ciphertext.
    '''        
    return _______________________________________________________________________
```

<!-- grade_id: ct, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-44-15b94d1fa268> in <module>
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: test-ct, points: 1, grade: true, solution: false -->
```python
# tests
key = 'ZEBRAS'
plaintext = 'WEAREDISCOVEREDFLEEATONCE'
ciphertext = 'EVLNACDTESEAROFODEECWIREE'
assert ct_encrypt(plaintext, key) == ciphertext
assert ct_decrypt(ciphertext, key) == plaintext
```

```python
import math
# key=input("Enter keyword text (Contains unique letters only): ").lower().replace(" ", "")
# plain_text = input("Enter plain text (Letters only): ").lower().replace(" ", "")
key = 'ZEBRAS'
plain_text = 'WEAREDISCOVEREDFLEEATONCE'
ciphertext = 'EVLNACDTESEAROFODEECWIREE'
```

```python
import math


```

Output:
```
[['X', 'X', 'X', 'X', 'X', 'X'], ['X', 'X', 'X', 'X', 'X', 'X'], ['X', 'X', 'X', 'X', 'X', 'X'], ['X', 'X', 'X', 'X', 'X', 'X'], ['X', 'X', 'X', 'X', 'X', 'X']]
```

```python

```

Output:
```
Encryption
Plain text is fz: WEAREDISCOVEREDFLEEATONCE
Cipher text is fz: EVLNXACDTXESEAXROFOXDEECXWIREE
```

```python
def ct_encrypt(plain_text, key):
    len_key = len(key)
    len_plain = len(plain_text)
    row = int(math.ceil(len_plain / len_key))
    matrix = [ ['X']*len_key for i in range(row) ]

#     print(matrix)
    t = 0
    len_key = len(key)
    
    for r in range(row):
        for c,ch in enumerate(plain_text[t : t+ len_key]):
            matrix[r][c] = ch
        t += len_key

    # print(matrix)
    sort_order = sorted([(ch,i) for i,ch in enumerate(key)])    #to make alphabetically order of chars
    # print(sort_order)

    cipher_text = ''
    for ch,c in sort_order:
        for r in range(row):
            cipher_text += matrix[r][c]

    print("Encryption")
    print("Plain text is fz:",plain_text)
    print("Cipher text is fz:",cipher_text)

key = 'ZEBRAS'
plain_text = 'WEAREDISCOVEREDFLEEATONCE'
ct_encrypt(plain_text, key)
```

Output:
```
Encryption
Plain text is fz: WEAREDISCOVEREDFLEEATONCE
Cipher text is fz: EVLNXACDTXESEAXROFOXDEECXWIREE
```

```python
def ct_decrypt(cipher_text, key):
    len_key = len(key)
    len_cipher_text = len(cipher_text)
    row = int(math.ceil(len_cipher_text / len_key))
    
    matrix_new = [ ['X']*len_key for i in range(row) ]
    key_order = [ key.index(ch) for ch in sorted(list(key))]    #to make original key order when we know keyword
    # print(key_order)

    t = 0
    for c in key_order:
        for r,ch in enumerate(cipher_text[t : t+ row]):
            matrix_new[r][c] = ch
        t += row
    # print(matrix_new) 

    p_text = ''
    for r in range(row):
        for c in range(len_key):
            p_text += matrix_new[r][c] if matrix_new[r][c] != 'X' else ''

    print("Decryption")
    print("Cipher text is:",cipher_text)
    print("Plain text is :",p_text)
    
key = 'ZEBRAS'
plain_text = 'WEAREDISCOVEREDFLEEATONCE'
cipher_text = 'EVLNACDTESEAROFODEECWIREE'

ct_decrypt(cipher_text, key)



```

Output:
```
Decryption
Cipher text is: EVLNACDTESEAROFODEECWIREE
Plain text is : ECOEWADDVIRTELROEENEFSCAE
```

```python
from itertools import*
def f(a,b):a,c,s=a.upper(),''.join,sorted;return c(map(c,zip_longest(*[v for k,v in s([(v,b[i::len({*a})])for i,v in enumerate(s({*a},key=a.find))])],fillvalue='')))
```

```python
plain_text
```

Output:
```
'WEAREDISCOVEREDFLEEATONCE'
```

```python
key
```

Output:
```
'ZEBRAS'
```

```python
f(cipher_text, key)
```

Output:
```
'ASZBRE'
```

```python
key = 'ZEBRAS'
plain_text = 'WEAREDISCOVEREDFLEEATONCE'
cipher_text = 'EVLNACDTESEAROFODEECWIREE'

encrypt(plain_text, key)
```

Output:
```
'EVLNACDTESEAROFODEECWIREE'
```

```python
def encrypt(message, keyword):
    matrix = createEncMatrix(len(keyword), message)
    keywordSequence = getKeywordSequence(keyword)

    ciphertext = "";
    for num in range(len(keywordSequence)):
        pos = keywordSequence.index(num+1)
        for row in range(len(matrix)):
            if len(matrix[row]) > pos:
                ciphertext += matrix[row][pos]
    return ciphertext


def createEncMatrix(width, message):
    r = 0
    c = 0
    matrix = [[]]
    for pos, ch in enumerate(message):
        matrix[r].append(ch)
        c += 1
        if c >= width:
            c = 0
            r += 1
            matrix.append([])

    return matrix


def getKeywordSequence(keyword):
    sequence = []
    for pos, ch in enumerate(keyword):
        previousLetters = keyword[:pos]
        newNumber = 1
        for previousPos, previousCh in enumerate(previousLetters):
            if previousCh > ch:
                sequence[previousPos] += 1
            else:
                newNumber += 1
        sequence.append(newNumber)
    return sequence

```

```python
def decrypt(message, keyword):
    matrix = createDecrMatrix(getKeywordSequence(keyword), message)

    plaintext = "";
    for r in range(len(matrix)):
        for c in range (len(matrix[r])):
            plaintext += matrix[r][c]
    return plaintext


def createDecrMatrix(keywordSequence, message):
    width = len(keywordSequence)
    height = len(message) / width
    if height * width < len(message):
        height += 1

    matrix = createEmptyMatrix(width, height, len(message))

    pos = 0
    for num in range(len(keywordSequence)):
        column = keywordSequence.index(num+1)

        r = 0
        while (r < len(matrix)) and (len(matrix[r]) > column):
            matrix[r][column] = message[pos]
            r += 1
            pos += 1

    return matrix


def createEmptyMatrix(width, height, length):
    matrix = []
    totalAdded = 0
    for r in range(height):
        matrix.append([])
        for c in range(width):
            if totalAdded >= length:
                return matrix
            matrix[r].append('')
            totalAdded += 1
    return matrix


def getKeywordSequence(keyword):
    sequence = []
    for pos, ch in enumerate(keyword):
        previousLetters = keyword[:pos]
        newNumber = 1
        for previousPos, previousCh in enumerate(previousLetters):
            if previousCh > ch:
                sequence[previousPos] += 1
            else:
                newNumber += 1
        sequence.append(newNumber)
    return sequence

```

```python
key = 'ZEBRAS'
plain_text = 'WEAREDISCOVEREDFLEEATONCE'
cipher_text = 'EVLNACDTESEAROFODEECWIREE'

decrypt(cipher_text, key)
```

Output:
```
<ipython-input-137-17445bcfa12f> in createDecrMatrix(keywordSequence, message)
     15         height += 1
     16 
---> 17     matrix = createEmptyMatrix(width, height, len(message))
     18 
     19     pos = 0

<ipython-input-137-17445bcfa12f> in createEmptyMatrix(width, height, length)
     33     matrix = []
     34     totalAdded = 0
---> 35     for r in range(height):
     36         matrix.append([])
     37         for c in range(width):

TypeError: 'float' object cannot be interpreted as an integer
```

```python
def argsort(seq):
    '''A helper function that returns the tuple of indices that would sort the
    sequence seq.'''
    return tuple(x[0] for x in sorted(enumerate(seq), key=lambda x: x[1]))


def ct_idx(length, key):
    '''A helper function that returns the tuple of indices that would permute 
    the letters of a message according to the key using the irregular case of 
    columnar transposition cipher.'''
    last_row_size = length % len(key)
    num_rows = length // len(key) + (1 if last_row_size else 0)
    seq = tuple(range(length))
    matrix = [[1,2,3],[4,5,6]]
    return tuple(matrix[i][j] for j in argsort(key)
                 for i in range(num_rows - (1 if j >= last_row_size else 0)))


def ct_encrypt(plaintext, key):
    '''
    Return the ciphertext of a plaintext by the key using the irregular case
    of columnar transposition cipher.

    Parameters
    ----------
    plaintext (str): a message in uppercase without punctuations/spaces.
    key (str): secret key to encrypt plaintext.
    '''
    return ''.join([plaintext[i] for i in ct_idx(len(plaintext), key)])


def ct_decrypt(ciphertext, key):
    '''
    Return the plaintext of the ciphertext by the key using the irregular case
    of columnar transposition cipher.

    Parameters
    ----------
    ciphertext (str): a string in uppercase without punctuations/spaces.
    key (str): secret key to decrypt ciphertext.
    '''        
    return 'hello'
```

```python
sorted(key)
```

Output:
```
['A', 'B', 'E', 'R', 'S', 'Z']
```

```python
# tests
key = 'ZEBRAS'
plaintext = 'WEAREDISCOVEREDFLEEATONCE'
ciphertext = 'EVLNACDTESEAROFODEECWIREE'

argsort(key) # ZEBRAS becomes sorted like this ->>> ['A', 'B', 'E', 'R', 'S', 'Z'] then to this ->>> (4, 2, 1, 3, 5, 0)
ct_encrypt(plaintext, key)
```

Output:
```
<ipython-input-84-e4bf4bdf8a37> in ct_idx(length, key)
     13     seq = tuple(range(length))
     14     matrix = [[1,2,3],[4,5,6]]
---> 15     return tuple(matrix[i][j] for j in argsort(key)
     16                  for i in range(num_rows - (1 if j >= last_row_size else 0)))
     17 

<ipython-input-84-e4bf4bdf8a37> in <genexpr>(.0)
     14     matrix = [[1,2,3],[4,5,6]]
     15     return tuple(matrix[i][j] for j in argsort(key)
---> 16                  for i in range(num_rows - (1 if j >= last_row_size else 0)))
     17 
     18 

IndexError: list index out of range
```